---
title       : Insert the chapter title here
description : Insert the chapter description here

--- type:NormalExercise lang:r xp:50 skills:1 key:cca61c83cf
## A


*** =instructions


*** =hint


*** =pre_exercise_code
```{r}
```

***= sample_code
```{r}
```

***= solution
```{r}
```

*** =sct
```{r}
success_msg("ok")
```
--- type:NormalExercise lang:r xp:50 skills:1 key:28eed07e35
## A


*** =instructions


*** =hint


*** =pre_exercise_code
```{r}
```

***= sample_code
```{r}
```

***= solution
```{r}
```

*** =sct
```{r}
success_msg("ok")
```
